# Assignment 1
For this assignment you are welcomed to use other regex resources such a regex "cheat sheets" you find on the web. Feel free to share good resources with your peers in slack!



Before start working on the problems, here is a small example to help you understand how to write your own answers. In short, the solution should be written within the function body given, and the final result should be returned. Then the autograder will try to call the function and validate your returned result accordingly. 

```python
def example_word_count():
    # This example question requires counting words in the example_string below.
    example_string = "Amy is 5 years old"
    
    # YOUR CODE HERE.
    # You should write your solution here, and return your result, you can comment out or delete the
    # NotImplementedError below.
    result = example_string.split(" ")
    return len(result)

    #raise NotImplementedError()
```

## Part A

Find a list of all of the names in the following string using regex.

<!-- grade_id: names, points: 0, grade: false, solution: true -->
```python
import re
def names():
    simple_string = """Amy is 5 years old, and her sister Mary is 2 years old. 
    Ruth and Peter, their parents, have 3 kids."""

    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: correct_names, points: 1, grade: true, solution: false -->
```python
assert len(names()) == 4, "There are four names in the simple_string"

```

## Part B

The dataset file in [assets/grades.txt](assets/grades.txt) contains a line separated list of people with their grade in 
a class. Create a regex to generate a list of just those students who received a B in the course.

<!-- grade_id: grades, points: 0, grade: false, solution: true -->
```python
import re
def grades():
    with open ("assets/grades.txt", "r") as file:
        grades = file.read()

    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: correct_grades, points: 1, grade: true, solution: false -->
```python
assert len(grades()) == 16

```

## Part C

Consider the standard web log file in [assets/logdata.txt](assets/logdata.txt). This file records the access a user makes when visiting a web page (like this one!). Each line of the log has the following items:
* a host (e.g., '146.204.224.152') 
* a user_name (e.g., 'feest6811' **note: sometimes the user name is missing! In this case, use '-' as the value for the username.**)
* the time a request was made (e.g., '21/Jun/2019:15:45:24 -0700')
* the post request type (e.g., 'POST /incentivize HTTP/1.1' **note: not everything is a POST!**)

Your task is to convert this into a list of dictionaries, where each dictionary looks like the following:
```
example_dict = {"host":"146.204.224.152", 
                "user_name":"feest6811", 
                "time":"21/Jun/2019:15:45:24 -0700",
                "request":"POST /incentivize HTTP/1.1"}
```

<!-- grade_id: logs, points: 0, grade: false, solution: true -->
```python
import re
def logs():
    with open("assets/logdata.txt", "r") as file:
        logdata = file.read()
    
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-correct_logs, points: 1, grade: true, solution: false -->
```python
assert len(logs()) == 979

one_item={'host': '146.204.224.152',
  'user_name': 'feest6811',
  'time': '21/Jun/2019:15:45:24 -0700',
  'request': 'POST /incentivize HTTP/1.1'}
assert one_item in logs(), "Sorry, this item should be in the log results, check your formating"

```

I would like to apply to your master thesis advertisement in the group research on the
 subject of "                 " in Volkswagen AG as a very successful and innovative solution provider in the field which offers many exciting positions.
 
For me, finding a puzzling pattern or discovering something novel in a row collection of data, is as exciting as a new scientific breakthrough. The unpredictable nature of ML and deep learning and their applications in AI, data science and computer vision, make them very exciting. I feel like an ancient explorer, finding a new land when trying to find hidden patterns in a large dataset. Accordingly, with relying on very good knowledge in mathematics and physics, gained within bachelor’s degree, I was confident to choose this filed as my major topic my master degree to pursue in my future career. As a result, with starting my master course, I began a very long learning journey. I have studied in-depth machine learning courses and a student job about CNNs at TU Chemnitz.at the same time I got the advantages of cutting age online courses.  I have been participating in many machine learning and deep learning projects. The building blocks of Deep NNs and CNNs, residual network (ResNet), Car detection with YOLO, implementation of a FaceNet with Keras, art generation with Neural Style Transfer, implementation of Recurrent Neural Network (RNN) to build a character-level text generation, trigger word detection are among them. 
I am doing my internship project in AUDI AG in the field of data science and ML in classification and clustering. During this project, I have been testing and analyzing data with both ML and statistical analyses. This program has given me the opportunity to examine my theoretical knowledge in a real industrial environment to gain more intuition in related subjects. 
in my opinion ML techniques are very general and can be applied to a wide range of problems with little domain-specific knowledge.

I believe that in this position, I will be in my way to achieve my goals. I can assure you will find me a responsible and reliable person with the ability of good teamwork and appropriate communication skills, which allows me to work effectively not only independently but also as a team player. I am quite confident in seeing myself in more significant and more complicated projects. I am very excited and motivated to learn more about possible opportunities at Volkswagen Company and share how I will be an excellent fit for your offering position(s). 
